<problem display_name="Matching" markdown="null" weight="0.0">
  <script type="text/javascript">
$(document).ready(function(){
	console.log('outer ready');
});

// Logs a javascript object.
function logThatThing(ThatThing){

	// Log it to the console just to verify it's working
	console.log(JSON.stringify(ThatThing));
	
	// Send it to the official edX logamajig!
	Logger.log("harvardx.template.matching_zoo", ThatThing);

}

</script>
<script type="loncapa/python">


from python_lib import HXGraders
right_answer = []

##################### Set These #########################
  
right_answer.append( [['A','2'],['B','1'],['C','3']] )
partial_credit = True
feedback = True
############## Leave Everything Else Alone ##############

def answercheck(e, ans):

  # This pulls from the python_lib.zip file in Files &amp; Uploads.
  return HXGraders.matchingAGrader(ans, right_answer, partial_credit, feedback)

 
</script>
  <p>This is a matching exercise. Drag the elements on the left into the ones on the right to create a match. Click on the <span style="color:darkblue">[-]</span> to remove the match. You can match items with more than one other item, but you should choose the best match if you can. Click on the individual items for a better view if you need it.</p>
  <p>This problem type gives partial credit! Each correct assignment is worth one point; each error takes a point away. Play around with different matches to see how it works.</p>
  <customresponse cfn="answercheck">
    <jsinput gradefn="MatchingA.getGrade" get_statefn="MatchingA.getState" set_statefn="MatchingA.setState" width="780" height="550" html_file="/static/MatchingA_Zoo.html" sop="false"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>Zebras are herbivores, lions are carnivores, and bears are omnivores. Setting "bear" as both "carnivore" and "herbivore" won't work, because those are exclusive categories - they say "only" or "exclusively."</p>
    </div>
  </solution>
</problem>
